---
title: Introducción
description: 'Conceptos básicos sobre modelización estadística y presentación de algunos ejemplos de modelos lineales.'
free_preview: true
---

##  

```yaml
type: VideoExercise
key: 86112b0c24
xp: 50
```

`@projector_key`
eda455040b30125f4beac71ed9193c29

---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
